```python
from grading import get_grade, submit_grade
```

The goal of the following exercises is to do a demo on the grading process!

## Exercise 1

Create a string variable `hello`  with value "hello".

<!-- grade_id: cell-04d8a63358035edf, points: 0, grade: false, solution: true -->
```python
# hello = 
### BEGIN SOLUTION
hello = "hello"
### END SOLUTION
```

<!-- grade_id: cell-544982d3bdbdc1ff, points: 10, grade: true, solution: false -->
```python
assert hello == "hello"
```

## Exercise 2

Create a string variable 'world'  with value "world".

<!-- grade_id: cell-e90f0cd8bd1cc3eb, points: 0, grade: false, solution: true -->
```python
#world = 
### BEGIN SOLUTION
world = "world"
### END SOLUTION
```

<!-- grade_id: cell-37346dd062259224, points: 10, grade: true, solution: false -->
```python
assert world == "world"
```

## Grading Process

congratulation, you completed SLU0! 🥳 
Let's grade it 💪 and send some feedback to the instructors.      

In order to do that, first, please save the notebook. For that, click on 💾, top bar of the notebook, first icon at left.   
Then, run the cell below and answer to the questions.

In order to register your account on the grading system, and save tour grade, you will need your slack Member ID.   
You can check your slack Member Id on your slack profile. 

Before starting, change the following variable `grading` from False to True.

<!-- grade_id: cell-c64ce8bc30e3a0f2, points: 0, grade: false, solution: true -->
```python
#grading = !replace me with True
slack_id = None  # example: "UTS63FC02"
slack_usename = None # example: "my slack usernamr"
email = None # example: "my_email@lisbondatascience.org"
notebook_saved = False #change to the value True after saving this notebook
### BEGIN SOLUTION
slack_id = "ADMINID"
slack_username = "ADMINUSERNAME"
email = "ADMINEMAIL"
notebook_saved = True
### END SOLUTION
```

```python
#run me
assert slack_id is not None, "Please replace this variable with your slack Member Id"
assert slack_username is not None, "Please replace this variable with your slack username"
assert email is not None, "Please replace this variable with your email"
assert notebook_saved is True, "Please make sure you save the notebook before running this cell and change this variable to True"
grade, max_grade = get_grade("Exercise Notebook.ipynb")
submit_grade(slack_id, slack_username, email, int(grade), max_grade, 0)
```

Output:
```
{"learning_unit": 0, "slack_id": "ADMINID", "grade": 10, "metadata": {"slack_username": "ADMINUSERNAME", "email": "ADMINEMAIL", "max_grade": 20.0}}
Whoopsie Daisy
```

```python

```