```python
from grading import login, send_grade
```

The goal of this notebook is to do a demo on the grading process!

## Exercise 1

Create a string variable `hello`  with value "hello".

<!-- grade_id: cell-04d8a63358035edf, points: 0, grade: false, solution: true -->
```python
# hello = 
### BEGIN SOLUTION
hello = "hello"
### END SOLUTION
```

<!-- grade_id: cell-544982d3bdbdc1ff, points: 10, grade: true, solution: false -->
```python
assert hello == "hello"
```

## Exercise 2

Create a string variable 'world'  with value "world".

<!-- grade_id: cell-e90f0cd8bd1cc3eb, points: 0, grade: false, solution: true -->
```python
#world = 
### BEGIN SOLUTION
world = "world"
### END SOLUTION
```

<!-- grade_id: cell-37346dd062259224, points: 10, grade: true, solution: false -->
```python
assert world == "world"
```

## Grading Process

congratulation, you completed SLU0! 🥳 
Let's grade it 💪 and send some feedback to the instructors.      

In order to do that, first, please save the notebook. For that, click on 💾, top bar of the notebook, first icon at left.   
Then, run the cell below and answer to the questions.

In order to register your account on the grading system, and save tour grade, you will need your slack Member ID.   
You can check your slack Member Id on your slack profile. 

Before starting, change the following variable `grading` from False to True.

<!-- grade_id: cell-c64ce8bc30e3a0f2, points: 0, grade: false, solution: true -->
```python
#!replace me with True
### BEGIN SOLUTION
grading = True
### END SOLUTION
grading = False
```

```python
#run me
if grading:
    while True:
        slack_member_id = login()
        if slack_member_id:
            grade = send_grade(slack_member_id, "slu0")
            if grade:
                print("Process Completed!")
                break
            else:
                print("Let's try again!")
                continue
        else:
            print("Let's try again!")
            continue
```