```python
from submit import submit
```

The goal of the following exercises is to do a demo on the grading process!

## Exercise 1

Create a string variable `hello`  with value "hello".

<!-- grade_id: cell-04d8a63358035edf, points: 0, grade: false, solution: true -->
```python
# hello = 
### BEGIN SOLUTION
hello = "hello"
### END SOLUTION
```

<!-- grade_id: cell-544982d3bdbdc1ff, points: 10, grade: true, solution: false -->
```python
assert hello == "hello"
```

## Exercise 2

Create a string variable 'world'  with value "world".

<!-- grade_id: cell-e90f0cd8bd1cc3eb, points: 0, grade: false, solution: true -->
```python
# world = 
### BEGIN SOLUTION
world = "world"
### END SOLUTION
```

<!-- grade_id: cell-37346dd062259224, points: 10, grade: true, solution: false -->
```python
assert world == "world"
```

# Sumit your work!

Congratulation, you completed SLU0! 🥳 
Let's grade it 💪 and send some feedback to the instructors.      

Before moving on, please run and save the current notebook. This step needed to insure that your grade is accordingly with the latest code.   
Then, follow the next step.

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and assign your slack id to a `slack_id` variable in the cell bellow.
Example:

```python
slack_id = "UTS63FC02"
```

<!-- grade_id: cell-c64ce8bc30e3a0f2, points: 0, grade: false, solution: true -->
```python
#slack_id = ...  
### BEGIN SOLUTION
slack_id = "ADMNSLU00"
### END SOLUTION
```

<!-- grade_id: cell-cdb84a6deb449648, points: 0, grade: true, solution: false -->
```python
#Run Me
slu = 0
### BEGIN HIDDEN TESTS
#Don't forget to change the slu number, it should be an int variable and correspondent to the number of the slu
### END HIDDEN TESTS
assert slack_id is not None
submit(slack_id, slu)
```

Output:
```
Your grade is 20 out of 20!
...Sending Grade...
Your grade was successfully submitted!
```